---
title: "Machine learning (1)"
author: "Guyliann Engels & Philippe Grosjean"
description: "**SDD III Module 1** Classification supervisée avec l'ADL."
tutorial:
  id: "C01Lb_ml1"
  version: 2.2.0/7
output:
  learnr::tutorial:
    progressive: true
    allow_skip: true
runtime: shiny_prerendered
---

```{r setup, include=FALSE}
BioDataScience3::learnr_setup()
SciViews::R("ml", "explore")

# exercice preparation --------
read("biometry", package = "BioDataScience") %>.%
  select(., gender, weight, height, wrist) %>.%
  drop_na(.) %->%
  bio

## Prepare train and test sets
set.seed(164)
bio_split <- initial_split(bio, prop = 2/3)
bio_train <- training(bio_split)
bio_test <- testing(bio_split)

## Create an LDA model
bio_lda <- mlLda(data = bio_train, gender ~ .)

## Confusion 
bio_conf <- confusion(predict(bio_lda, bio_test), bio_test$gender)
bio_conf_tab <- summary(bio_conf)
```

```{r, echo=FALSE}
BioDataScience3::learnr_banner()
```

```{r, context="server"}
BioDataScience3::learnr_server(input, output, session)
```

------------------------------------------------------------------------

## Objectifs

De nombreux algorithmes de classification supervisée coexistent. Nous découvrons ici comment classer des items à l'aide d'un ordinateur grâce à l'Analyse Discriminante Linéaire (ADL). Cette méthode se base sur les mêmes principes que l'ACP que vous avez utilisée au cours SDD II.

L'élaboration d'un classifieur requiert de suivre une série d'étapes :

-   Explorer les données
    -   Visualiser les données à l'aide de divers graphiques,
    -   Vérifier et corriger le type de chaque variable (surtout pour les variables "factor"),
    -   Calculer des variables supplémentaires pertinentes,
    -   Traiter le problème des valeurs manquantes si votre modèle ne les gère pas,
    -   Éliminer les variables inutiles (par exemple, le code d'identification d'un individu).
-   Préparer le set d'apprentissage et le set de test
-   Entraîner un classifieur (choisir l'algorithme, l'ajuster sur le set d'apprentissage...)
-   Évaluer la qualité du classifieur à l'aide des métriques pertinentes
-   Déployer le classifieur afin de prédire les classes dans un gros jeu de données non encore classé manuellement.

L'objectif de ce tutoriel est de vous permettre de découvrir l'analyse discriminante linéaire. Vous devez être bien conscient de toutes les étapes ci-dessus et vous devrez les mettre en œuvre dans le cadre de vos projets plus tard. Cependant, pour rendre cette première approche de la classification supervisée plus abordable, nous ne détaillerons pas toutes ces étapes, mais seulement les quelques unes qui mènent à un classifieur ADL fonctionnel.

## Différences morphométriques entre hommes et femmes

Le calcul de l'indice de masse corporelle (IMC) a déjà fait l'objet de plusieurs de vos analyses dans les cours précédents. C'est en effet un sujet à la fois simple à comprendre au premier abord, et en réalité, assez compliqué dans le détail. Le projet de départ est de définir un indice, un nombre, qui quantifie l'"obésité" d'une personne sur une échelle allant de l'anorexie la plus extrême jusqu'à l'obésité morbide la plus forte. On se rend vite compte qu'il ne suffit pas simplement de prendre la masse, ou même, la masse rapportée à la taille. En effet, un certain nombre de facteurs interviennent comme le genre, l'âge, la morphologie générale, la pratique de sport à haut niveau, etc.

À partir du moment où la classification supervisée permet de déterminer s'il est possible de différencier des sous-populations sur base de mesures quantitatives, nous pouvons utiliser l'ADL pour répondre à la question suivante : si les différences entre genres interviennent de manière non négligeables sur les variables qui sont utilisées régulièrement dans les formules de calcul de l'IMC, il doit être possible de *prédire* si un individu est un homme ou une femme *uniquement* sur base de l'information contenue dans ces variables quantitatives. Autrement dit, nous prenons le problème à l'envers. Au lieu de nous demander comment le genre influence l'IMC, on se demande si les variables quantitatives qui interviennent dans le calcul de l'IMC permettent de prédire le genre. Et si c'est le cas, nous en conclurons que le genre ne peut probablement pas être négligé dans nos formules. Cette approche inversée est utile dans certaines situations où les relations entre les variables sont obscures car il existe des algorithmes de classification supervisée qui fonctionnent comme des boîtes noires et pour lesquelles ce type d'information n'est justement pas nécessaire.

![](images/man_woman.png){width="35%"}

Vous avez à votre disposition le tableau de données suivant que l'on réduit à quatre variables :

-   le genre `gender`, que l'on cherchera à prédire, une variable `factor` à deux niveaux `M` pour homme et `W` pour femme, et
-   trois variables quantitatives rencontrées dans les diverses formules d'IMC, soit la masse `weight` en kg, la taille `height`, et la circonférence du poignet `wrist` en mm.

```{r, echo=TRUE}
read("biometry", package = "BioDataScience",lang = "fr") %>.%
  sselect(., gender, weight, height, wrist) %->%
  bio
```

Commençons par explorer ce jeu de données.

```{r, echo=TRUE}
skimr::skim(bio)
```

Étant donné que la LDA se base sur les mêmes principes que l'ACP, une ACP peut-être réalisée pour résumer les données.

```{r, echo=TRUE}
bio_red <- sdrop_na(bio)
bio_pca <- pca(data = sselect(bio_red, -gender), ~ ., scale = TRUE)
summary(bio_pca)
a <- chart(bio_pca, type = "loadings")
b <- chart(bio_pca, type = "scores", labels = bio_red$gender) +
  stat_ellipse()
combine_charts(list(a, b))
```

```{r qu_explo}
question("Sélectionnez les éléments pertinents de l'analyse exploratoire réalisée ci-dessus.",
  answer("Le tableau de données comprend quatre variables numériques et une variable facteur ordonnée."),
  answer("La répartition entre les hommes (M) et les femmes (W) est proche de 50/50", correct = TRUE),
  answer("On observe la présence de valeurs manquantes pour la variable de la circomférence du poignets.", correct = TRUE),
  answer("Toutes les variables sont complètes. Il n'y a pas de valeurs manquantes dans ce tableau."),
  answer("L'ACP permet de séparer en grande partie les hommes des femmes. Les hommes ont tendance à être plus grands et plus lourds.", correct = TRUE),
  allow_retry = TRUE, random_answer_order = TRUE,
  correct = "Vous avez sélectionné les éléments corrects.",
  incorrect = "Attention, Ce n'est pas complet. Intéressez-vous avec plus d'attention à la description du tableau et à l'ACP.",
  submit_button = "Soumettre une réponse",
  try_again_button = "Resoumettre une réponse")
```

## Préparation des données

Suite à la phase exploratoire de l'analyse, il en est ressorti que des valeurs manquantes devaient être filtrées.

```{r, echo=TRUE}
bio <- sdrop_na(bio, wrist) # Filtrer les valeurs manquantes en spécifiant la colonne concernée
```

Utilisez les fonctions `initial_split()`, `training()` et `testing()` pour définir votre set d'apprentissage et votre set de test. Votre set d'apprentissage va se nommer `bio_train` et votre set de test sera `bio_test`. Il vous est demandé de réaliser un set d'apprentissage contenant 3/4 des observations et en conservant les mêmes proportions qu'au départ pour le genre.

```{r split_h2, exercise=TRUE}
set.seed(164) # Fixer le début du générateur de nombres pseudo-aléatoires
bio_split <- initial_split(___, prop = ___, ___ = gender)
bio_split
# Diviser le tableau
___ <- training(___)
bio_test <- testing(___)
```

```{r split_h2-hint-1}
set.seed(164) # Fixer le début du générateur de nombres pseudo-aléatoires
bio_split <- initial_split(bio, prop = 3/4, ___ = gender)
bio_split
# Diviser le tableau
___ <- training(___)
bio_test <- testing(___)

## Attention, le prochain indice est la solution ##
```

```{r split_h2-solution}
## Solution ##
set.seed(164) # Fixer le début du générateur de nombres pseudo-aléatoires
bio_split <- initial_split(bio, prop = 3/4, strata = gender)
bio_split
# Diviser le tableau
bio_train <- training(bio_split)
bio_test <- testing(bio_split)
```

```{r split_h2-check}
grade_code("Vous avez réalisez votre set d'apprentissage et votre set de test avec les proportions demandées. N'hésitez pas à consulter la page d'aide des fonctions `initial_split()`, `training()` et `testing()`. Vous y trouverez des informations très intéressantes sur la réalisation des sets d'apprentissage et de test.", "Avez-vous bien respecté les consignes ? Il ne faut compléter que les zones signalées par `___`.")
```

## Entrainement du modèle

Créez maintenant un classifieur du genre en fonction des autres variables (en utilisant une formule abrégée). Choisissez le tableau adéquat pour ce faire et placer ce classifieur dans la variable `bio_lda`.

```{r lda1_h2, exercise=TRUE}
bio_lda <- ___(data = ___, ___ ~ ___)
summary(bio_lda)
```

```{r lda1_h2-hint-1}
bio_lda <- mlLda(data = ___, gender ~ ___)
summary(bio_lda)

## Attention, le prochain indice est la solution ##
```

```{r lda1_h2-solution}
## Solution ##
bio_lda <- mlLda(data = bio_train, gender ~ .)
summary(bio_lda)
```

```{r lda1_h2-check}
grade_code("Voici donc votre premier classifieur LDA. Voyons maintenant quoi en faire...", "Avez-vous bien écrit la formule sous sa forme condensée ?")
```

## Performances du classifieur

Vous allez maintenant déterminer les performances de votre classifieur `bio_lda`. Ne vous trompez pas dans les jeux de données à utiliser `bio`, `bio_train` et `bio_test` pour ce faire.

```{r lda2_h2, exercise=TRUE}
bio_pred <- predict(___, ___)
bio_conf <- confusion(___, ___$___)
bio_conf
summary(bio_conf)
```

```{r lda2_h2-hint-1}
bio_pred <- predict(bio_lda, ___)
bio_conf <- confusion(bio_pred, ___$___)
bio_conf
summary(bio_conf)

## Attention, le prochain indice est la solution ##
```

```{r lda2_h2-solution}
## Solution ##
bio_pred <- predict(bio_lda, bio_test)
bio_conf <- confusion(bio_pred, bio_test$gender)
bio_conf
summary(bio_conf)
```

```{r lda2_h2-check}
grade_code("Vous venez de réaliser vos premières analyses des performances d'un classifieur. Prenez un peu de temps pour analyser votre matrice de confusion et les métriques qui en découlent et répondez aux questions suivantes.")
```

```{r qu_lda}
quiz(
  question("Combien d'items du set de test sont correctement classé ?",
    answer(sprintf("%1.f", sum(bio_conf_tab$TP)), correct = TRUE),
    answer(sprintf("%1.f", sum(bio_conf_tab$Auto))),
    answer(sprintf("%1.f", bio_conf_tab$Manu[1])),
    answer(sprintf("%1.f", bio_conf_tab$TN[2])),
    answer("Aucune des réponses proposées"),
    allow_retry = TRUE, random_answer_order = TRUE,
    submit_button = "Soumettre une réponse",
    try_again_button = "Resoumettre une réponse",
    incorrect = "Mauvaise réponse. Concentrez-vous et réessayez.",
    correct = "C'est du bon travail."),
  question("Quel est le taux d'erreur globale (en %) ?",
    answer(sprintf("%.1f", 100 * (1 - (sum(bio_conf_tab$TP)/sum(bio_conf_tab$Auto)))), correct = TRUE),
    answer(sprintf("%.1f", 100 * (sum(bio_conf_tab$TP)/sum(bio_conf_tab$Auto)))),
    answer(sprintf("%3.f", sum(bio_conf_tab$Auto) - sum(bio_conf_tab$TP))),
    answer(sprintf("%3.f", bio_conf_tab$TN[2])),
    answer("Aucune des réponses proposées"),
    allow_retry = TRUE, random_answer_order = TRUE,
    submit_button = "Soumettre une réponse",
    try_again_button = "Resoumettre une réponse",
    incorrect = "Mauvaise réponse. Réessayez après avoir mieux examiné les résultats obtenus.",
    correct = "C'est correct."),
  question("Quel est le taux de vrais positifs pour les hommes (M) ?",
    answer(sprintf("%.3f", bio_conf_tab[row.names(bio_conf_tab) == "M", ]$Recall), correct = TRUE),
    answer(sprintf("%.3f", bio_conf_tab[row.names(bio_conf_tab) == "M", ]$Fscore)),
    answer(sprintf("%3.f", sum(bio_conf_tab$Auto) - sum(bio_conf_tab$TP))),
    answer(sprintf("%.3f", bio_conf_tab[row.names(bio_conf_tab) == "W", ]$Recall)),
    answer("Aucune des réponses proposées"),
    allow_retry = TRUE, random_answer_order = TRUE,
    submit_button = "Soumettre une réponse",
    try_again_button = "Resoumettre une réponse",
    incorrect = "Mauvaise réponse. Examinez bien la matrice de confusion.",
    correct = "C'est exact.")
  )
```

## Conclusion

Vous venez de créer votre premier classifieur. Nous avons détaillé ensemble les étapes d'exploration et de préparation des données rapidement. Ensuite, vous avez été guidé dans la séparation en set d'apprentissage et de test, et dans l'utilisation du premier pour entraîner votre LDA et du second pour déterminer les performances du classifieur ainsi obtenu.

À l'issue de cette étude, nous constatons que dans ce jeu de données, il est possible de différencier les hommes des femmes à l'aide seulement des informations de masse, taille et circonférence de poignet avec un taux d'erreur inférieur à 15%. Nous en concluons donc que l'information portée par ces trois variables quantitatives permet cette discrimination de genre, et donc que nous devons nous en préoccuper lors de l'élaboration de notre prochain indice de masse corporel qui mettra tout le monde d'accord !

```{r comm_noscore, echo=FALSE}
question_text(
  "Laissez-nous vos impressions sur cet outil pédagogique",
  answer("", TRUE, message = "Pas de commentaires... C'est bien aussi."),
  incorrect = "Vos commentaires sont enregistrés.",
  placeholder = "Entrez vos commentaires ici...",
  allow_retry = TRUE
)
```
