---
title: "Découverte de learnr"
author: "Guyliann Engels & Philippe Grosjean"
description: "**SDD III** Quelques exercices simples pour découvrir learnr."
tutorial:
  id: "C00La_decouverte"
  version: 2.0.0/5
output:
  learnr::tutorial:
    progressive: true
    allow_skip: true
runtime: shiny_prerendered
---
```{r setup, include=FALSE}
BioDataScience3::learnr_setup()
SciViews::R()
```

```{r, echo=FALSE}
BioDataScience3::learnr_banner()
```

```{r, context="server"}
BioDataScience3::learnr_server(input, output, session)
```

----

## Préambule

Vous vous trouvez au sein d'un tutoriel interactif construit avec `learnr`. La première chose à vérifier à l’ouverture du tutoriel interactif est votre `nom d’utilisateur` et votre `adresse email`. En effet, votre progression sera enregistrée, mais cela ne peut se faire que si vous renseignez ces données correctement avant de travailler dans le tutoriel learnr.

Le learnr est un outil pédagogique mis au point afin de proposer des tutoriels interactifs comprenant des illustrations, des questions à choix multiples, des exercices R, ...

Des questions ouvertes (sous la forme zones de code R) vous sont proposées dans les exercices. Elles vous permettent d’expérimenter directement des instructions dans R depuis le document learnr. Pour exécuter ces instructions, il faut cliquer sur `Run Code`. Vous pouvez le faire autant de fois que vous le voulez. Modifiez le code, cliquez `Run Code`, analysez le résultat, modifiez votre code, recliquez `Run Code`, etc… jusqu’à ce que vous soyez satisfait du résultat. Finissez l’exercice et soumettez votre réponse en cliquant sur le bouton `Submit Answer`. **Tant que vous n'aurez pas cliqué sur `Submit Answer`, votre exercice restera inachevé et ne comptera donc pas dans l'évaluation de votre progression.**

Si vous n'avez jamais utilisé de tutoriel learnr, familiarisez-vous d'abord avec son interface [ici](https://wp.sciviews.org/sdd-umons-2020/learnr.html).


## Objectifs

- Découvrir les quiz de type `learnr` composés de questions ouvertes et de questions à choix multiples.
- Bien comprendre la logique d'utilisation des zones de code R.

## Questions à choix multiples

Notre planète terre est en train de subir des grands changements climatiques en partie liés à l'activité humaine. Entre autres, elle a tendance à se réchauffer.

```{r quiz_chauffe}
quiz(
  question("Quel gaz dans l'atmosphère est considéré comme le principal responsable de cette augmentation de température par effet de serre ?",
    answer("N$_2$"),
    answer("O$_2$"),
    answer("O$_3$"),
    answer("CO$_2$", correct = TRUE),
    answer("SO$_3$"),
    answer("NH$_3$"),
    allow_retry = TRUE,
    incorrect = "Mauvaise réponse. Recommencez afin de trouver la bonne réponse, c'est la combustion d'énergie fossile qui est le mécanisme le plus incriminé... ça devrait vous aider.",
    correct = "Bravo, c'est correct ! En effet, la concentration en CO$_2$ a très fortement augmenté dans l'atmosphère suite à la combustion d'énergie fossile en masse par les humains depuis le début de l'ère indistrielle."
  ),
  question("Quelle est l'année la plus chaude (température moyenne annuelle à Uccle) en Belgique entre 1833 et 2018 ?",
    answer("1989", message = "Avec une moyenne annuelle de 11,26°C, l'année 1989 se place à la 8e position"),
    answer("2007", message = "Avec une moyenne annuelle de 11,52°C, l'année 2007 se place à la 4e position"),
    answer("2014", correct = TRUE),
    answer("2018", message = "Avec une moyenne annuelle de 11,85°C, l'année 2018 se place à la 2e position"),
    allow_retry = TRUE,
    incorrect = "Non, désolé ! Recommencez afin de trouver la bonne réponse.",
    correct = "Bravo, c'est correct ! L'année 2014 est l'année la plus chaude
      enregistrée entre 1833 et 2018 avec une température moyenne annuelle de 11,93°C. 
      Depuis 1999 toutes les années font partie du top 20 à  l'exception de 2010, 2012 et 2013."),
  question("Les changements climatiques sont souvent associés à l'**écologie**. Quel(s) célèbre(s) scientifique(s) ont contribué à cette discipline ?",
    answer("Ernst Haeckel", correct = TRUE, message = "Ernst Haeckel est même l'inventeur du terme écologie."),
    answer("Charles Darwin", message = "Charles Darwin a écrit l'**Origine des espèces** qui établit les premiers pas de la théorie de l'évolution."),
    answer("Carl von Linné", message = "Carl von Linné propose le système de classification binomiale des espèces."),
    answer("Alexander von Humboldt", correct = TRUE, message = "Alexander von Humboldt n'a pas inventé 
      le terme d'écologie. Il est cependant, le père de l'écologie scientifique en
      trouvant des associations entre la température et les espèces, par exemple."),
    random_answer_order = TRUE,
    allow_retry = TRUE,
    incorrect = "Incorrect, recommencez afin de trouver la bonne réponse.",
    correct = "Bravo, c'est correct ! Ernst Haeckel et Alexander von Humboldt sont deux précurseur de l'écologie.")
)
```

Comme vous avez pu vous en rendre compte en répondant aux questions ci-dessus, elles peuvent être recommencées afin de trouver la bonne réponse. De plus, les réponses bonnes comme mauvaises peuvent être associées à une explication complémentaire. Elles servent essentiellement à votre **auto-évaluation** pour que vous puissiez vérifier que vous avez bien compris les concepts abordés.

En pratique dans le cours, il est fortement conseillé de revoir la matière immédiatement si vous constatez que vous n'arrivez pas à répondre aux questions... sinon vous risquez d'accumuler un retard qui pourrait devenir irrattrapable à la longue. Si vous trichez ici, la seule personne qui en payera le prix, c'est vous\ : cliquer et recommencer au hasard sans réfléchir est une *très mauvaise stratégie* qui ne vous donnera que l'illusion d'avancer. Dans ce cas, vous n'arriverez pas à faire les exercices plus compliqués qui arriveront ensuite, et vous finirez tout de même en échec\ !

## Questions ouvertes

Répondez aux questions ouvertes suivantes\ : 

La température moyenne annuelle en surface de la Mer du Nord est de 13.5°C en 2020. D'après les dernières recherches la température de cette mer augmente de 0.053°C annuellement.

- Quelle sera la température moyenne en 2030\ ?

*Utilisez la zone de code R ci-dessous pour répondre à la question. Vous pouvez l'utiliser comme une calculatrice.*

*Afin de tester vos instructions, cliquez sur `Run Code` (autant que vous voulez). Lorsque vous êtes satisfait de votre réponse, cliquez sur `Submit Answer`. Si l'exercice est trop compliqué, n'hésitez pas à utiliser les indices mis à votre disposition (`Hints`) s'il y en a, comme ici.*

```{r sea_temp1_h3, exercise=TRUE}

```

```{r sea_temp1_h3-hint-1}
# Vous pouvez utiliser la zone de code R comme une calculatrice.
# Addition : +
# Soustraction : - 
# Multiplication : *
# Division : /
```

```{r sea_temp1_h3-hint-2}
 # Additionner le nombre de degrés pris par année
 # à la temp&rature de 2020
 
 #### ATTENTION: Hint suivant = solution !####
```

```{r sea_temp1_h3-solution}
## Solution ##
13.5 + 0.053 * (2030 - 2020) 
```

```{r sea_temp1_h3-check}
grade_result(
  pass_if(~ identical(.result, 13.5 + (0.053 * (2030 - 2020))), "Déjà un demi degré en plus. Vous avez compris le principe de la zone de code R pour y effectuer vos calculs."),
  fail_if(~ TRUE, "Pas exactement. Pour chaque année écoulée, vous devez ajouter 0.053 degrés.")
)
```

- Quelle sera la température moyenne en 2100 (utilisez une formulation qui vous permettra de faire facilement le calcul pour n'importe quelle autre année avec un changement minimal, c'est-à-dire, juste l'année désirée)\ ?

*Quelquefois, un bouton offrant la solution est proposé, comme ci-dessous. N'utilisez-le qu'en dernier recours, si vous n'arrivez vraiment pas à solutionner le problème. Certains exercices demanderont que vous écriviez votre code R d'une façon bien précise. C'est le cas ici.*

```{r sea_temp2, exercise=TRUE, exercices.lines=5}
t0 = 13.5
y0 = 2020
y1 = 2100
deltat = 0.053
___ + ___ * (___ - ___)
```

```{r sea_temp2-solution}
## Solution ##
t0 = 13.5
y0 = 2020
y1 = 2100
deltat = 0.053
t0 + deltat * (y1 - y0) 
```

```{r sea_temp2-check}
grade_code("Vos début en code R sont impressionnants.")
```

Comme vous pouvez le constater, la réponse de l'exercice peut être donnée comme dernier indice de la question (bouton `Hints`) ou via le bouton `Solution`. **Encore une fois, autant pour les suggestions que pour les solutions, utiliser ces aides sans réfléchir juste pour répondre à la question en trichant ne trompera que vous\ : vous serez toujours incapable de répondre réellement à ce genre de question... et donc incapable de résoudre les exercices plus difficiles qui suivront\ !**

Si vous avez été incapable de répondre à la question sans aller jusqu'à visualiser la réponse, essayez de décortiquer cette réponse pour déterminer ce que vous n'avez pas bien compris. Revoyez la matière correspondante du cours, discutez-en avec vos collègues, ou posez des questions à vos enseignants.

## Conclusion

Bravo\! Vous venez de terminer votre séance d'exercices dans un tutoriel `learnr`. Rappelez-vous qu'il s'agit d'un outil d'**auto-évaluation** mis à votre disposition pour vérifier que vous avez bien compris et assimilé les concepts de base avant de passer à des exercices plus difficiles. Utilisez-le de manière responsable.

```{r quiz_eval_noscore}
question("Comment évalueriez-vous globalement cet outils pédagogiques ?",
  answer("Excellent", correct = TRUE),
  answer("Bon", correct = TRUE),
  answer("Satisfaisant", correct = TRUE),
  answer("Insuffisant", correct = TRUE),
  answer("Très insuffisant", correct = TRUE),
  allow_retry = TRUE,
  type = "single", correct = "Merci, votre avis est enregistré.")
```

Le travail déjà réalisé dans un learnr est mémorisé d'une session à l'autre. Dans certains cas, il se peut cependant que cette mémoire de l'état d'avancement se perde. C'est le cas si vous changer d'ordinateur ou d'explorateur web. Mais cela se produit aussi si une nouvelle version du learnr est installée. Dans ce cas, *votre activité reste enregistrée* au niveau du rapport de progression (mais vérifiez toujours par vous-même, et contactez vos enseignants en cas de doute)\ !

```{r comm_noscore, echo=FALSE}
question_text(
  "Laissez-nous vos impressions sur cet outil pédagogique",
  answer("", TRUE, message = "Pas de commentaires... C'est bien aussi."),
  incorrect = "Vos commentaires sont enregistrés.",
  placeholder = "Entrez vos commentaires ici...",
  allow_retry = TRUE
)
```

Si vous voulez, vous pouvez encore tester d'autres commandes R ci-dessous\ :

```{r code_noscore, exercise=TRUE, exercise.lines=8, exercise.checker=learndown::checker_ack_learnr}
# Un commentaire est toujours précédé d'un '#' en R
```

```{r code_noscore-check}

```
